---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 400e0729fd
xp: 50
```

`@projector_key`
5b24923d6aba7f6de6db165e9e40b397

---

## Introduction to Logistic Regression

```yaml
type: VideoExercise
key: 3eed02ee7a
xp: 50
```

`@projector_key`
972fe35ea4e257689c1ab53913553465
